## Problem 3: How long distance individuals have travelled? (8 points)

In this problem the aim is to calculate the distance in meters that the individuals have travelled according the social media posts (Euclidean distances between points). In this problem, we will need the `userid` -column an the points created in the previous problem. You will need the shapefile `Kruger_posts.shp` generated in Problem 2 as input file.

Our goal is to answer these questions based on the input data:

 - What was the shortest distance travelled in meters?
 - What was the mean distance travelled in meters?
 - What was the maximum distance travelled in meters?

**In your code, you should first:**
 - Import required modules
 - Read in the shapefile as a geodataframe called `data`

<!-- grade_id: read_data, points: 0, grade: false, solution: true -->
```python
import os
import pandas as pd
import geopandas as gpd
import matplotlib.pyplot as plt
from shapely.geometry import Point, LineString
from pyproj import CRS

data = gpd.read_file('Kruger_posts.shp')


```

 - Check the crs of the input data. If this information is missing, set it as epsg:4326 (WGS84).
 - Reproject the data from WGS84 to `EPSG:32735` -projection which stands for UTM Zone 35S (UTM zone for South Africa) to transform the data into metric system. (don't create a new variable, update the existing variable `data`!)

<!-- grade_id: reproject, points: 0, grade: false, solution: true -->
```python
data = gpd.GeoDataFrame(data, geometry='geometry', crs=CRS.from_epsg(4326).to_wkt())

```

<!-- grade_id: problem_3_gdp_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
print(data.head())
```

Output:
```
lat        lon         timestamp    userid  \
0 -24.980792  31.484633  2015-07-07 03:02  66487960   
1 -25.499225  31.508906  2015-07-07 03:18  65281761   
2 -24.342578  30.930866  2015-03-07 03:38  90916112   
3 -24.854614  31.519718  2015-10-07 05:04  37959089   
4 -24.921069  31.520836  2015-10-07 05:19  27793716   

                     geometry  
0  POINT (-24.98079 31.48463)  
1  POINT (-25.49922 31.50891)  
2  POINT (-24.34258 30.93087)  
3  POINT (-24.85461 31.51972)  
4  POINT (-24.92107 31.52084)
```

<!-- grade_id: problem_3_CRS_test, points: 1.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
# Check that the crs is correct after re-projecting (should be epsg:32735)
print(data.crs)
```

Output:
```
GEOGCRS["WGS 84",DATUM["World Geodetic System 1984",ELLIPSOID["WGS 84",6378137,298.257223563,LENGTHUNIT["metre",1]]],PRIMEM["Greenwich",0,ANGLEUNIT["degree",0.0174532925199433]],CS[ellipsoidal,2],AXIS["geodetic latitude (Lat)",north,ORDER[1],ANGLEUNIT["degree",0.0174532925199433]],AXIS["geodetic longitude (Lon)",east,ORDER[2],ANGLEUNIT["degree",0.0174532925199433]],USAGE[SCOPE["unknown"],AREA["World"],BBOX[-90,-180,90,180]],ID["EPSG",4326]]
```

 - Group the data by userid

<!-- grade_id: group_data, points: 0, grade: false, solution: true -->
```python
grouped = data.groupby('userid')
grouped
# line3 = LineString(data['geometry'].unique())

# for key, group in grouped:
#     print('La Key', key)
# #     print('El Group', group)
#     group_geom = grouped['geometry']
#     print(group_geom)
#     for elem in group_geom:
#         print('el Elem es:', elem)
```

Output:
```
<pandas.core.groupby.generic.DataFrameGroupBy object at 0x7f301ce72490>
```

<!-- grade_id: problem_3_grouped_test, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
assert len(grouped.groups) == data["userid"].nunique(), "Number of groups should match number of unique users!"
```

**Then:**
- Create an empty GeoDataFrame called `movements`
- Create a for-loop where you iterate over the grouped object. For each user's data: 
    - [sort](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.sort_values.html) the rows by timestamp 
    - create a LineString object based on the user's points
    - add the geometry and the userid into the `movements` dataframe (one userid per row). You can achieve this either by using the `.at` indexer, or the `append` method. See hints for more help.
- Set the CRS of the ``movements`` GeoDataFrame as ``EPSG:32735`` 

<!-- grade_id: create_lines, points: 0, grade: false, solution: true -->
```python
movements = gpd.GeoDataFrame()

for group in grouped:
    
#     data.sort_values(by= ['timestamp'])
#     line = LineString(([lat, lon]) for lat, lon in zip(data['lat'], data['lon']))                                                                 
    
#     grouped.at[0, 'geometry']= movements
    movements = movements.append({'geometry':group})

    

```

Output:
```
<ipython-input-100-957584ffcff3> in <module>
      7 
      8 #     grouped.at[0, 'geometry']= movements
----> 9     movements = movements.append({'geometry':group})
     10 
     11 

/usr/lib/python3/dist-packages/pandas/core/frame.py in append(self, other, ignore_index, verify_integrity, sort)
   7095                 other = Series(other)
   7096             if other.name is None and not ignore_index:
-> 7097                 raise TypeError(
   7098                     "Can only append a Series if ignore_index=True"
   7099                     " or if the Series has a name"

TypeError: Can only append a Series if ignore_index=True or if the Series has a name
```

```python
for geometry in zip(grouped.geometry):
    geometry.apply(lambda x: LineString(x.tolist()))
#     print(geo_df.grouped.geometry.apply(lambda x: LineString(x.tolist())))
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-37-d256041e28a7> in <module>
      1 for geometry in zip(grouped.geometry):
----> 2     geometry.apply(lambda x: LineString(x.tolist()))
      3 #     print(geo_df.grouped.geometry.apply(lambda x: LineString(x.tolist())))

AttributeError: 'tuple' object has no attribute 'apply'
```

<!-- grade_id: problem_3_movements_tests, points: 3, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
movements.head()
```

**Finally:**
- Check once the crs definition of your dataframe (should be epsg:32735, define the correct crs if this information is missing)
- Calculate the lenghts of the lines into a new column called ``distance`` in ``movements`` GeoDataFrame.

<!-- grade_id: calculate_distance, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_distances_test, points: 2, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
movements.head()
```

You should now be able to print answers to the following questions: 

 - What was the shortest distance travelled in meters?
 - What was the mean distance travelled in meters?
 - What was the maximum distance travelled in meters?

```python

```

- Finally, save the movements of into a Shapefile called ``some_movements.shp``

<!-- grade_id: save_file, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_result_tests, points: 0.5, grade: true, solution: false -->
```python
# NON-EDITABLE CODE CELL FOR TESTING YOUR SOLUTION
import os
assert os.path.isfile(fp), "output shapefile does not exits"
```

That's all for this week!